# Aula 00 - Python Básico

# Exercício 01:

Crie uma função que, dado dois textos, retorne uma __lista__ (que pode ser vazia) em ordem lexicográfica com as palavras que ocorrem em ambos. A função não deve ser case sensitive, nesse caso, transforme todo o texto reescrevendo as palavras em letra minúscula. Para testar sua função, você pode utilizar os dois poemas abaixo, cuja saída esperada é: ['and', 'are', 'is'] 

_The rose is a rose,  
And was always a rose.  
But the theory now goes  
That the apple’s a rose,  
And the pear is, and so’s  
The plum, I suppose.  
The dear only knows  
What will next prove a rose.  
You, of course, are a rose –  
But were always a rose._  

_Roses are red,  
Violets are blue,  
sugar is sweet,  
And so are you._  



__Dicas:__  
não se preocupe em remover a pontuação do texto.  
a função sorted(lista) ordena em ordem crescente (ou lexicográfica) uma lista qualquer.  
a função string.lower() tranforma todos os caracteres da string em letra minúscula.  


<!-- grade_id: cell1-a00, points: 0, grade: false, solution: true -->
```python
#Exemplo:
poem1 = """The rose is a rose,
And was always a rose.
But the theory now goes
That the apple’s a rose,
And the pear is, and so’s
The plum, I suppose.
The dear only knows
What will next prove a rose.
You, of course, are a rose –
But were always a rose."""

poem2 = """Roses are red,
Violets are blue,
sugar is sweet,
And so are you."""


def common_words(text1, text2):
    
    some_set = sorted({x for x in text1.split() for y in text2.split() if x == y and x is not '\n'})
    some_list = [x.lower() for x in some_set]
    return some_list

```

<!-- grade_id: cell1t1-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
"""Check that common_words returns the correct output for several inputs"""

poem1 = """The rose is a rose,
And was always a rose.
But the theory now goes
That the apple’s a rose,
And the pear is, and so’s
The plum, I suppose.
The dear only knows
What will next prove a rose.
You, of course, are a rose –
But were always a rose."""

poem2 = """Roses are red,
Violets are blue,
sugar is sweet,
And so are you."""

assert_equal(common_words(poem1, poem2), ['and', 'are', 'is'])
```

<!-- grade_id: cell1t2-a00, points: 1, grade: true, solution: false -->
```python

```

# Exercício 02:

Crie uma função que, dado uma lista de tuplas representando pares de palavras com sua respectiva traduação e uma frase, retorne a frase traduzida. As palavras que não possuem uma tradução devem ser impressas como aparecem na frase de entrada. Para testar sua função, você pode utilizar a entrada abaixo, cuja saída esperada é: o cara tossiu.

Lista de tuplas: [("galaxy", "cara tossiu"), ("kagayaku", "canalha do"), ("atsuki", "alto que"), ("yuushi", "util")]  
Frase: "o galaxy"


__Observação:__  
Esse problema, conhecido como O Fantástico Jaspion! foi adaptado de https://br.spoj.com/problems/JASPION/.


<!-- grade_id: cell2-a00, points: 0, grade: false, solution: true -->
```python
#Exemplo:
lista_de_tuplas = [("galaxy", "cara tossiu"), ("kagayaku", "canalha do"), ("atsuki", "alto que"), ("yuushi", "util")]
frase = "o galaxy"


def jaspion(tuples_list, sentence):
    
    list_words = sentence.split()
    finalSentence = ""
    
    for word in list_words:
        for tupl in tuples_list:
            if word in tupl:
                if word == tupl[0]:
                    finalSentence += tupl[1] + " "
                elif word == tupl[1]: 
                    finalSentence += tupl[0] + " "
                break
        else:
            finalSentence += word + " "
    
    return finalSentence.strip()
    
    

```

<!-- grade_id: cell2t1-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
"""Check that jaspion returns the correct output for several inputs"""

tl1 = [("galaxy", "cara tossiu"), ("kagayaku", "canalha do"), ("atsuki", "alto que"), ("yuushi", "util")]
s1 = "o galaxy"

assert_equal(jaspion(tl1, s1), "o cara tossiu")
```

<!-- grade_id: cell2t2-a00, points: 1, grade: true, solution: false -->
```python

```