# Aula 05 - Probabilidade

# Exercício 01:


Suponha que a altura de mulheres adultas de algumas regiões seguem uma distribuição normal com $\mu = 162$ centímetros e $\sigma = 8$. Nesse caso, responda às perguntas abaixo:


(a) Dado que uma mulher mede $180$ centímetros, qual a probabilidade de alguém escolhido ao acaso ser maior que ela? Para responder à questão, crie uma função _a()_, sem parâmetros, que retorna a resposta da questão com uma precisão de 4 casas decimais.  

__Dica:__  
a função _round(var, n)_ retorna o valor da variável var com uma precisão de n casas decimais.

<!-- grade_id: cell11-a00, points: 0, grade: false, solution: true -->
```python
# Crie aqui a função a() - com esse nome e sem parâmetros - 
# para retornar a resposta com precisão de 4 casas decimais!
import math

# YOUR CODE HERE
def normal_pdf(x, mu=0, sigma=1):
    sqrt_two_pi = math.sqrt(2 * math.pi)
    return (math.exp(-(x-mu) ** 2 / 2 / sigma ** 2) / (sqrt_two_pi * sigma))

def a():
    print('oi')
    result = 0
    return round(result, 4)
```

<!-- grade_id: cell11t-a00, points: 1, grade: true, solution: false -->
```python

```

(b) Uma treinadora dessa região quer montar uma equipe de basquete. Para isso, ela quer delimitar uma altura mínima $h$ que as jogadoras devem ter. Ele quer que $h$ seja maior que pelo menos $90\%$ das alturas de mulheres daquela região. Qual o valor de $h$? Para responder à questão, crie uma função _b()_, sem parâmetros, que retorna a resposta da questão com uma precisão de 4 casas decimais.  

__Dica:__  
a função _round(var, n)_ retorna o valor da variável var com uma precisão de n casas decimais.

<!-- grade_id: cell12-a00, points: 0, grade: false, solution: true -->
```python
#Crie aqui a função b() - com esse nome e sem parâmetros - 
# para retornar a resposta com precisão de 4 casas decimais!

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell12t-a00, points: 1, grade: true, solution: false -->
```python

```

# Exercício 02:


As seguintes amostras foram geradas seguindo uma distribuição normal N($\mu$, $\sigma$), onde $\mu$, $\sigma$ não necessariamente são os mesmos para ambas. Nos histogramas gerados é possível visualizar essa distribuição.


```python
matplotlib inline
```

```python
from matplotlib import pyplot as plt

dados1 = [3.8739066,4.4360658,3.0235970,6.1573843,3.7793704,3.6493491,7.2910457,3.7489513,5.9306145,5.3897872,
          5.9091607,5.2491517,7.1163771,4.1930465,-0.1994626,3.2583011,5.9229948,1.8548338,4.8335581,5.2329008,
          1.5683191,5.8756518,3.4215138,4.7900996,5.9530234,4.4550699,3.3868535,5.3060581,4.2124300,7.0123823,
          4.9790184,2.2368825,3.9182012,5.4449732,5.7594690,5.4159924,3.5914275,3.4382886,4.0706780,6.9489863,
          6.3269462,2.8740986,7.4210664,4.6413206,4.2209699,4.2009752,6.2509627,4.9137823,4.9171593,6.3367493]

dados2 = [2.291049832,5.092164483,3.287501109,4.152289011,4.534256822,5.513028947,2.696660244,3.270482741,
          5.435338467,6.244110011,1.363583509,5.385855994,6.069527998,2.148361858,6.471584096,4.953202949,
          6.827787432,4.695468536,2.047598339,8.858080081,5.436394723,7.849470791,4.053545595,3.204185038,
          2.400954454,-0.002092845,3.571868529,6.202897955,5.224842718,4.958476608,6.708545254 -0.115002497,
          5.106492712,3.343396551,5.984204841,3.552744920,4.041155327,5.709103288,3.137316917,2.100906915,
          4.379147487,0.536031040,4.777440348,5.610527663,3.802506385,3.484180306,7.316861806,2.965851553,
          3.640560731,4.765175164,7.047545215,5.683723446,5.048988000,6.891720033,3.619091771,8.396155189,
          5.317492252,2.376071049,4.383045321,7.386186468,6.554626718,5.020433071,3.577328839,5.534419417,
          3.600534876,2.172314745,4.632719037,4.361328042,4.292156420,1.102889101,4.621840612,4.946746104,
          6.182937650,5.415993589,4.346608293,2.896446739,3.516568382,6.972384719,3.233811405,4.048606672,
          1.663547342,4.607297335 -0.753490459,3.205353052,1.269307121,0.962428478,4.718627886,4.686076530,
          2.919118501,6.204058666,4.803050149,4.670632749,2.811395731,7.214950058,3.275492976,2.336357937,
          8.494097155,6.473022507,8.525715511,4.364707111]

plt.hist(dados1)
plt.show()
plt.hist(dados2)
plt.show()
```

__a)__ A partir dos histogramas, tente aproximar uma normal a cada um deles, desenhando-a sobre o histograma. Para isso, você deve estimar valores de $\mu$ e $\sigma$. Não se esqueça de normalizar os dados, ou seja, o eixo y deve estar um uma escala de 0 a (no máximo) 1!

<!-- grade_id: cell21-a00, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

__b)__ Dado as duas amostras e as respectivas normais estimadas no exercício anterior, verifique através do plot ___qq-plot___ se é razoável pensar que as amostras vieram dessas distribuições. Além disso, compare o qq-plot das duas amostras, verificando por exemplo a relação entre o tamanho da amostra e estimação da distribuição utilizada para gerá-la. Utilize a última célula desse exercício para escrever sobre.



__Dicas:__  
Mais informações sobre o qq-plot em: https://www.itl.nist.gov/div898/handbook/eda/section3/qqplot.htm  

<!-- grade_id: cell22-a00, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell23-a00, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Exercício 03:


Dado uma tabela com informações sobre uma amostra com 20 alunos contendo a nota desses alunos em algumas disciplinas e os níveis de dificuldade das mesmas, crie uma função que retorne a probabilidade condicional estimada à partir dos dados para dois eventos dados, informando ainda se os eventos são independentes ou não. Ou seja, dado a tabela mostrada no exemplo (lista de listas) e dois eventos A e B, retorne a probabilidade condicional de A dado B (P(A|B)) com uma precisão de 4 casas decimais. O retorno da função, entretanto, deve ser uma frase (string) escrita da seguinte forma: _str: val_ onde _str_ é a string "Independentes" se os eventos A e B são independentes e "Dependentes" caso contrário e _val_ é o valor da probabilidade condicional P(A|B) com uma precisão de 4 casas decimais.


__Dica:__  
a função format(var, '.nf') retorna uma string com o valor da variável var com uma precisão de exatamente n casas decimais.


<!-- grade_id: cell3-a00, points: 0, grade: false, solution: true -->
```python
# Esses dados se referem às notas (A-E) de 20 alunos de acordo com a dificuldade da disciplina (Fácil ou Difícil) 
# Coluna 1: id do aluno
# Coluna 2: dificuldade da disciplina ('Facil' ou 'Dificil')
# Coluna 3: nota do aluno (A-E)
data = [[1, 'Facil', 'C'],
        [2, 'Facil', 'A'],
        [3, 'Dificil', 'E'],
        [4, 'Dificil', 'B'],
        [5, 'Dificil', 'B'],
        [6, 'Dificil', 'A'],
        [7, 'Facil', 'D'],
        [8, 'Dificil', 'C'],
        [9, 'Facil', 'D'],
        [10, 'Facil', 'C'],
        [11, 'Facil', 'A'],
        [12, 'Facil', 'A'],
        [13, 'Dificil', 'B'],
        [14, 'Dificil', 'C'],
        [15, 'Dificil', 'E'],
        [16, 'Dificil', 'C'],
        [17, 'Facil', 'A'],
        [18, 'Dificil', 'D'],
        [19, 'Facil', 'B'],
        [20, 'Facil', 'A']]


def prob_cond(tabela, A, B):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell3t1-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal
"""Check that prob_cond returns the correct output for several inputs"""

assert_equal(prob_cond(data, 'A', 'Facil'), 'Dependentes: 0.5000')
assert_equal(prob_cond(data, 'E', 'Facil'), 'Dependentes: 0.0000')
assert_equal(prob_cond(data, 'A', 'Dificil'), 'Dependentes: 0.1000')
assert_equal(prob_cond(data, 'E', 'Dificil'), 'Dependentes: 0.2000')
```

<!-- grade_id: cell3t2-a00, points: 1, grade: true, solution: false -->
```python

```