# Aula 04 - Estatística

# Exercício 01:


Utilizando os dados de acidentes fatais em companhias aéreas dos Estados Unidos de 1985 a 1999, calcule algumas estatísticas básicas. Você deve retornar uma __lista__ com os valores das estatísticas calculadas, sendo elas, nessa ordem: menor valor, maior valor, média, mediana, dispersão, variância, desvio-padrão e intervalo-interquartil. Para responder à questão, crie uma função _estat(acidentes)_ que retorna a lista com os valores correspondentes às resposta da questão, inteiros quando forem inteiros ou com uma precisão de 4 casas decimais caso contrário.

__Teste:__
`assert_equal(estat(acidentes), ans)`, sendo que `ans` é uma lista contendo os valores corretos para as estatísticas que este exercício pede.

__Dicas:__  
1) A função round(var, n) retorna o valor da variável var com uma precisão de n casas decimais.

2) Execute o teste `assert_equal(estat(lista_boba), ans_bobo)` para alguma `lista_boba` que você saiba calcular as estatísticas no papel.


__Fonte:__ https://aviation-safety.net/

<!-- grade_id: cell1-a00, points: 0, grade: false, solution: true -->
```python
# Crie aqui a função estat(acidentes) - com esse nome e parâmetro - 
# a função deve retornar a lista com as respostas com precisão de 4 casas decimais!

acidentes = [0,14,0,1,0,4,1,0,0,2,1,5,0,3,0,0,6,1,1,12,3,1,5,0,3,0,0,1,1,0,1,5,2,1,1,3,4,0,3,0,2,2,1,0,1,1,1,3,0,4,3,8,7,3,0,1]
teste = [1,2,3,4,5,6,7,8,9,10]

# YOUR CODE HERE
from nose.tools import assert_equal
import numpy as np

def estat(acidentes):
    p = 4
    menor_valor = min(acidentes)
    maior_valor = max(acidentes)
    media  = np.mean(acidentes)
    mediana = np.median(acidentes)
    dispersao = maior_valor - menor_valor
    variancia  = np.var(acidentes)
    desvio_padrao = np.std(acidentes)
    intervalo_interquartil = np.subtract(sorted(acidentes)[int(0.75*len(acidentes))], sorted(acidentes)[int(0.25*len(acidentes))])
    
    lista = [menor_valor, maior_valor, media, mediana, dispersao, variancia, desvio_padrao, intervalo_interquartil]
    estatistica = [round(x, p) for x in lista]
    return estatistica

assert_equal(estat(teste), [1, 10, 5.5000,5.5000, 9, 8.2500, 2.8723, 5])
```

<!-- grade_id: cell1t-a00, points: 1, grade: true, solution: false -->
```python

```

# Exercício 02:


Procure encontrar correlações espúrias interessantes e apresente um exemplo encontrado. Ou seja, aprensente dois conjuntos de dados que possuem alta correlação (muito positivas ou muito negativas) sem que um seja de fato o causador do outro. Além disso, deixe resgistrado os gráficos com a distribuição dos dados e um gráfico de dispersão como forma de visualizar a correlação entre os dados. Calcule a covariância e correlação entre os dados e, por fim, se possível, tente explicar qual poderia ser a verdadeira causa da ocorrência das observações. Para isso, utilize a última célula desse notebook.


__Observação:__  
Para ideias de correlações espúrias, veja os seguintes sites:  
http://tylervigen.com/spurious-correlations  
https://en.wikipedia.org/wiki/Spurious_relationship#Other_relationships  


<!-- grade_id: cell21-a00, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell22-a00, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

# Exercício 03:

O Paradoxo de Simpson também pode ser observado na taxa de sobrevivência do [naufrágio do RMS Titanic](https://en.wikipedia.org/wiki/RMS_Titanic). Havia uma estimativa de 2.224 passageiros e tripulantes a bordo, e mais de 1.500 morreram, tornando-se um dos mais mortais desastres marítimos comerciais em tempos de paz da história moderna. Dados sobre os passageiros podem ser baixados [aqui](https://ww2.amstat.org/publications/jse/v3n3/datasets.dawson.html). Este conjunto de dados foi processado por mim e transformado em um arquivo [csv](https://www.dropbox.com/s/vk8jf0wyczqxkvv/survival_titanic.csv?dl=0), que é muito mais fácil de tratar que um arquivo texto. Informações sobre ele podem ser lidas [aqui](https://www.dropbox.com/s/xpjw74khyqx9ww4/survival_titanic.README.txt?dl=0). 

Embora muitas outras informações existam sobre os passageiros, aqui vamos trabalhar com apenas quatro:

1) Se o passageiro sobreviveu;  
2) A classe do seu bilhete (primeira, segunda ou terceira) ou se ele era membro da tripulação;  
3) O sexo do passageiro;  
4) E se ele era um adulto ou uma criança.  

Os dados estão organizados da seguinte maneira:

|Coluna|Descrição|Valores|
| :-: | :-: | :-: |
|0|Classe| 0 = tripulação, 1 = primeira, 2 = segunda, 3 = terceira|
|1|Idade|1 = adulto, 0 = criança|
|2|Sexo|1 = masculino, 0 = feminino|
|3|Sobreviveu?|1 = sim, 0 = não|


O seguinte trecho de código carrega os dados que serão usados para responder as questões a seguir:


```python
import pandas as pd
df = pd.read_csv('survival_titanic.csv')
X = df.values
```

__a)__ Verifique qual classe possui uma taxa de sobrevivência maior, para isso, implemente uma função _survivalRatePerClass(X, classe)_ com esse mesmo nome e parêmetros que, dado a tabela X e um número (0,1,2,3) representando a classe, retorna a proporção de pessoas daquela classe que sobreviveram com uma precisão de 5 casas decimais. 

__Dica:__  
a função _round(var, n)_ retorna o valor da variável var com uma precisão de n casas decimais.

<!-- grade_id: cell31-a00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

def survivalRatePerClass(X, classe):
    X_C = [p for p in X if p[0] == classe]
    rate = sum(p[3] for p in X_C)/len(X_C)
    return round(rate, 5)

for i in range(0, 4):
    print(i, survivalRatePerClass(X, i))
```

Output:
```
0 0.23955
1 0.62462
2 0.41404
3 0.25212
```

<!-- grade_id: cell3t1-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(survivalRatePerClass(X, 0), 0.23955)
```

<!-- grade_id: cell3t2-a00, points: 1, grade: true, solution: false -->
```python

```

__b)__ Verifique qual classe possui uma taxa de sobrevivência maior de acordo com a idade e sexo dos passageiros. Para isso, implemente uma função _survivalRatePerClassAndColumn(X, classe, coluna)_ com esse mesmo nome e parêmetros que, dado a tabela X, um número (0,1,2,3) representando a classe e um número (1,2) representando respectivamente as colunas que representam a idade e o sexo dos passageiros, retorna a proporção de pessoas daquela classe que sobreviveram com uma precisão de 5 casas decimais. Nesse caso, retorne uma tupla do tipo (<proporção de sobreviventes crianças da classe>, <proporção de sobreviventes adultos da classe>) quando a coluna for a 1 e (<proporção de sobreviventes mulheres da classe>, <proporção de sobreviventes homens da classe>)  quando a coluna for a 2.

__Dica:__  
a função _round(var, n)_ retorna o valor da variável var com uma precisão de n casas decimais.

<!-- grade_id: cell32-a00, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def survivalRatePerClassAndColumn(X, classe, coluna):
    X_C0 = [p for p in X if p[0] == classe and p[coluna]==0]
    X_C1 = [p for p in X if p[0] == classe and p[coluna]==1]
    sur_0, sur_1 = (None,None)
    rSur_0, rSur_1 = (None, None)
    if(X_C0):
        sur_0 = sum(p[3] for p in X_C0)/len(X_C0)
        rSur_0 = round(sur_0, 5)
    if(X_C1):
        sur_1 = sum(p[3] for p in X_C1)/len(X_C1)
        rSur_1 = round(sur_1, 5)
        
    return rSur_0, rSur_1

print(survivalRatePerClassAndColumn(X, 0, 1))
print(survivalRatePerClassAndColumn(X, 1, 1))
print(survivalRatePerClassAndColumn(X, 0, 2))
```

Output:
```
(None, 0.23955)
(1.0, 0.61755)
(0.86957, 0.22274)
```

<!-- grade_id: cell3t3-a00, points: 1, grade: true, solution: false -->
```python
from nose.tools import assert_equal

#classe 0 (criança, adulto)
assert_equal(survivalRatePerClassAndColumn(X, 0, 1), (None, 0.23955))
#classe 1 (criança, adulto)
assert_equal(survivalRatePerClassAndColumn(X, 1, 1), (1.0, 0.61755))
#classe 0 (mulher, homem)
assert_equal(survivalRatePerClassAndColumn(X, 0, 2), (0.86957, 0.22274))
```

<!-- grade_id: cell3t4-a00, points: 1, grade: true, solution: false -->
```python

```

__c)__ Após brincar com as funções criadas, faça uma breve análise sobre os resultados encontrados, relacionando-os ao Paradoxo de Simpson visto em sala de aula.

<!-- grade_id: cell33-a00, points: 1, grade: true, solution: true -->
Os resultados encontrados nos mostram que só verificar um tipo de estatística nos dados não é o suficiente. O Paradoxo de Simpson fica explicitado quando olhamos apenas para a função survivalRatePerClass(), que averigua a sobrevivência por classe, mas sem olhar a porcentagem de sobrevivência dentre grupos menores, como homens, mulheres e crianças (que é feita olhando as colunas separadamente), e sem considerar a proporção destes grupos por entre as classes.